# I AM NOT DONE

a_integer = 1
a_float = 0.1
a_string = "Hello World!"
a_bytes = b"Hello"
a_boolean = True
a_tuple = (1, 2, "3", False)

assert isinstance(a_integer, int)
assert isinstance(a_float, float)
assert isinstance(a_string, str)
assert isinstance(a_bytes, bytes)
assert isinstance(a_boolean, True)
assert isinstance(a_tuple, tuple)
