# Python is a language is strong and dynamic typing
# This means that the interpreter will dynamically set the variable type during attribuitions
# But once typed, the variable will have a type, respecting the protocols and behaviors of their type.

# Once you are done, remove the line below:
# I AM NOT DONE

x = 1
y = 0.1
z = 'Hello'

assert isinstance(x, int)
assert isinstance(y, float)
assert isinstance(z, str)
